Adult neurogenesis through glial transdifferentiation in a CNS injury paradigm

Author response:

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1:

Summary:

Casas-Tinto et al. present convincing data that injury of the adult Drosophila CNS triggers transdifferentiation of glial cells and even the generation of neurons from glial cells. This observation opens up the possibility of getting a handle on the molecular basis of neuronal and glial generation in the vertebrate CNS after traumatic injury caused by Stroke or Crush injury. The authors use an array of sophisticated tools to follow the development of glial cells at the injury site in very young and mature adults. The results in mature adults revealing a remarkable plasticity in the fly CNS and dispels the notion that repair after injury may be only possible in nerve cords which are still developing. The observation of so-called VC cells which do not express the glial marker repo could point to the generation of neurons by former glial cells.

Conclusion:

The authors present an interesting story that is technically sound and could form the basis for an in-depth analysis of the molecular mechanism driving repair after brain injury in Drosophila and vertebrates.

Strengths:

The evidence for transdifferentiation of glial cells is convincing. In addition, the injury to the adult CNS shows an inherent plasticity of the mature ventral nerve cord which is unexpected.

Weaknesses:

Traumatic brain injury in Drosophila has been previously reported to trigger mitosis of glial cells and generation of neural stem cells in the larval CNS and the adult brain hemispheres. Therefore this report adds to but does not significantly change our current understanding. The origin and identity of VC cells is unclear.

The Reviewer correctly points out that it has been reported that traumatic brain injury trigger generation of neural stem cells. However, according to previous reports, those cells where quiescent Dpn+ neuroblast. We now report that already differentiated adult neuropil glia transdifferentiate into neurons. Which is a new mechanism not previously reported.

We agree with the reviewer regarding the identity of VC neurons although according to the results of G-TRACE experiments the origin is clear, they originate from neuropil glia (i.e. Astrocyte-like glia and ensheathing glia). We have used a battery of antibodies previously reported to identify specific subtypes of neurons to identify these newly generated neurons (Figure S1). We did not find any other neuronal marker rather than Elav that co-localize with VC cells

Reviewer #2:

Summary:

Casas-Tinto et al., provide new insight into glial plasticity using a crush injury paradigm in the ventral nerve cord (VNC) of adult Drosophila. The authors find that both astrocyte-like glia (ALG) and ensheating glia (EG) divide under homeostatic conditions in the adult VNC and identify ALG as the glial population that specifically ramps up proliferation in response to injury, whereas the number of EGs decreases following the insult. Using lineagetracing tools, the authors interestingly observe the interconversion of glial subtypes, especially of EGs into ALGs, which occurs independent of injury and is dependent on the availability of the transcription factor Prospero in EGs, adding to the plasticity observed in the system. Finally, when tracing the progeny of differentiated glia, Casas-Tinto and colleagues detect cells of neuronal identity and provide evidence that such glia-derived neurogenesis is specifically favored following ventral nerve cord injury, which puts forward a remarkable way in which glia can respond to neuronal damage.

Numerous experiments have been carried out in 7-day-old flies, showing that the observed plasticity is not due to residual developmental remodeling or a still immature VNC.

By elegantly combining different genetic tools, the authors show glial divisions with mitotic-dependent tracing and find that the number of generated glia is refined by apoptosis later on.

The work identifies Prospero in glia as an important coordinator of glial cell fate, from development to the adult context, which draws further attention to the upstream regulatory mechanisms.

We express our gratitude to the reviewer for their keen appreciation of our efforts and their enthusiasm for the outcomes of this research.

Weaknesses:

Although the authors do use a variety of methods to show glial proliferation, the EdU data (Figure 1B) could be more informative (Figure 1B) by displaying images of non-injured animals and providing quantifications or the mention of these numbers based on results previously acquired in the system.

We appreciate the Reviewer’s comment. We believed that adding images of non-injured animals did not add new information as we already quantified the increase of glial proliferation upon injury in Losada-Perez let al. 2021. Besides, the purpose of this experiment was to figure out if dividing cells where Astrocyte-like glia rather than the number of dividing cells. Comparing independent experiments could be tricky but if we compare the quantifications of G2-M glia (repo>fly-Fucci) done in Losada-Perez et al 2021 (fig 1C) with the quantifications of G2-M neuropil glia done in this work (fig 1C) we can see that the numbers are comparable.

The experiments relying on the FUCCI cell cycle reporter suggested considerable baseline proliferation for EGs and ALGs, but when using an independent method (Twin Spot MARCM), mitotic marking was only detected for ALGs. This discrepancy could be addressed by assessing the co-localization of the different glia subsets using the identified driver lines with mitotic markers such as PH3.

In our understanding this discrepancy could be explained by the magnitude of proliferation. The lower proliferation rate of EG (as indicate the fly-fucci experiments) combining with the incomplete efficiency of MARCM clones induction reduces considerably the chances of finding EG MARCM clones. PH3 is a mitotic marker but it is also found in apoptotic cells (Kim and Park 2012. DOI: 10.1371/journal.pone.0044307) however, we stained injured VNCs with anti-Ph3 and found ALG cells positive for PH3 (Author response image 1).

The data in Figure 1C would be more convincing in combination with images of the FUCCI Reporter as it can provide further information on the location and proportion of glia that enter the cell cycle versus the fraction that remains quiescent.

We added a Figure 1 V2 (version 2) with the suggested images (1-C’).

The analyses of inter-glia conversion in Figure 3 are complicated by the fact that Prospero RNAi is both used to suppress EG - to ALG conversion and as a marker to establish ALG nature. Clarifications if the GFP+ cells still expressed Pros or were classified as NP-like GFP cells are required here.

As described in the text, Pros is a marker for ALG and the results suggest that Prospero expression is required for the EG to ALG transition. We clarified these concepts in the text accordingly. In figure 3 we showed images of NP-like cells originated from EG that are prospero+, and therefore supporting the transdifferentiation from EG to ALG.

The conclusion that ALG and EG glial cells can give rise to cells of neuronal lineage is based on glial lineage information (GFP+ cells from glial G-trace) and staining for the neuronal marker Elav. The use of other neuronal markers apart from Elav or morphological features would provide a more compelling case that GFP+ cells are mature neurons.

We completely agree with the reviewer's observation regarding the identity of VC neurons. We have used a battery of antibodies previously reported to identify specific subtypes of neurons to identify these newly generated neurons (Figure S1). We did not find any other neuronal marker rather than Elav that colocalize with VC cells

Although the text discusses in which contexts, glial plasticity is observed or increased upon injury, the figures are less clear regarding this aspect. A more systematic comparison of injured VNCs versus homeostatic conditions, combined with clear labelling of the injury area would facilitate the understanding of the panels.

We appreciate the Reviewer’s observation. We have carefully checked all figures and labelled then as “Injured” or “Not Injured”. We added a Figure 2-V2 and a figure 4-V2.

Context/Discussion

The study finds that glia in the ventral cord of flies have latent neurogenic potential. Such observations have not been made regarding glia in the fly brain, where injury is reported to drive glial divisions or the proliferation of undifferentiated progenitor cells with neurogenic potential.

Discussing this different strategy for cell replacement adopted by glia in the VNC and pointing out differences to other modes seems fascinating. Highlighting differences in the reactiveness of glia in the VNC compared to the brain also seems highly relevant as they may point to different properties to repair damage.

Based on the assays employed, the study points to a significant amount of

glial "identity" changes or interconversions, which is surprising under homeostatic conditions. The significance of this "baseline" plasticity remains undiscussed, although glia unarguably show extensive adaptations during nervous system development.

It would be interesting to know if the "interconversion" of glia is determined by the needs in the tissue or would shift in the context of selective ablation/suppression of a glial type.

We deeply appreciate the Reviewer’s enthusiasm on this subject, it is indeed fascinating. We made a reduced discussion in order to fit in the eLife Short report requirements but the specific condition that trigger glial interconversion are of great interest for us. To compromise EG or ALG viability and evaluate the behaviour of glial cells is of great interest for developmental biology and regeneration, but the precise scenario to develop these experiments is not well defined. In this report, we aim to reproduce an injury in Drosophila brain and this model should serve to analyze cellular behaviours. The scenario where we deplete on specific subpopulation of glial cells is conceptually attractive, but far away from the scope of this report.

Reviewer #3:

In this manuscript, Casas-Tintó et al. explore the role of glial cells in the response to a neurodegenerative injury in the adult brain. They used Drosophila melanogaster as a model organism and found that glial cells are able to generate new neurons through the mechanism of transdifferentiation in response to injury.

This paper provides a new mechanism in regeneration and gives an understanding of the role of glial cells in the process.

Changes from the earlier version to the revised version

Abstract

Abstract

As the global population ages, the prevalence of neurodegenerative disorders is fast increasing. This neurodegeneration as well as other CNS injuries cause permanent disabilities. Thus, generation of new neurons is the rosetta stone in contemporary neuroscience.

Glial cells support central nervous system (CNS) homeostasis through evolutionary conserved mechanisms. Upon damage, glial cells activate an immune and inflammatory response to clear the injury site from debris, and proliferate to restore cell number. This glial regenerative response (GRR) is mediated by the neuropil associated glia (NG) in Drosophila, equivalent to vertebrate astrocytes, oligodendrocytes (OL) and oligodendrocyte progenitor cells (OPCs). Here, we examine the contribution of NG lineages and the GRR in response to injury. The results indicate that NG exchanges identities between EG and ALG. Additionally, we found that NG cells undergo transdifferentiation to yield neurons. Moreover, this transdifferentiation increases in injury conditions. Thus, these data demonstrate that glial cells are able to generate new neurons through direct transdifferentiation. The present work makes a fundamental contribution to the CNS regeneration field and describes a new physiological mechanism to generate new neurons.

Introduction

Neuronal loss caused by injury or neurodegenerative diseases is irreversible in the Central Nervous System (CNS) of vertebrates. Spinal fractures represent 5-6% of all fractures, and spinal cord damage complicates 16-40% of spinal fractures causing significant impairment to patients (Ding et al., 2022). Moreover, according to the United Nations, the prevalence of neurodegenerative diseases will almost double in the next 20 years. Current therapeutic strategies aim to replace damaged or lost cells through dedifferentiation and transplantation or cell reprogramming of glial cells. Thus, understanding glial cell contribution after CNS damage is relevant to attain its functional recovery (Li et al., 2017).

Glial cells activate inflammatory and proliferative responses upon injury in a process known as Glial Regenerative Response (GRR) (Kato et al., 2018; Losada-Perez, 2018). This response is described in multiple species of vertebrates and invertebrates, and the molecular mechanisms underlying this process are conserved from Drosophila to mice or humans (Chandra et al., 2023; Kato et al., 2018, 2015, 2011; Li et al., 2020; Losada-Perez, 2018; Losada-Pérez et al., 2021; Purice et al., 2017; Smith et al., 1987). Glial proliferation is a well-known cellular response upon CNS damage (Chandra et al., 2023; Kato et al., 2015, 2011; Li et al., 2020; Li and Hidalgo, 2020; Purice et al., 2017) however it is not clear how this response contributes to CNS regeneration. In mammals, oligodendrocyte progenitor cells (OPCs or NG2-glia) react giving rise to oligodendrocytes while in Drosophila, neuropil associated glia play that role (Losada-Pérez et al., 2021). Upon injury, mammalian OPCs can produce astrocytes and oligodendrocytes (OL) contributing to CNS repair (Kato et al., 2018; Niu et al., 2021). GRR of neuropil glia (NPG) in Drosophila larva is modulated by a genetic network that controls the switch between quiescence, proliferation and differentiation. Although, whether these mechanisms operate in the mature differentiated adult brain is still unknown. In Drosophila, NPG somata reside in the cortex/neuropil interface (Figure 1A). NPG includes two groups: Astrocyte-Like Glia (ALG), and Ensheathing Glia (EG). EG enwraps the neuropil, and bundles of axons within the neuropil that define different functional regions (Losada-Perez, 2018). By contrast, ALG cells project to the neuropil forming a dense meshwork with the synapses (Losada-Perez, 2018), which suggests that ALG cells may play a significant role in maintaining the architecture and function of the neuropil. Larval ALG upregulates kon-tiki expression upon injury, a gene which is required for glial division (Losada-Perez et al., 2016). Besides, larval EG during metamorphosis can de-differentiate, proliferate and re-differentiate into adult ALG and EG (Kato et al., 2020). Thus, NPG appear as good candidates to replace lost cells with functional ones.

Neuropil glia proliferation in normal conditions and upon crush injury in adults

Earlier version: A, Schematic representation of Drosophila adult Central Nervous System (CNS) and Ventral Nerve Cord (VNC), indicating the MtN, where the injury is performed. B, Representative confocal image showing dividing cells (EdU+) with ALG identity (Repo+ and Pros+) and glial identity different from ALG (Repo+, Pros-). C, Number of ALG or EG cells in G2-M per MtN, in normal conditions and 24 h after injury. Paired t-test, *P<0.05, **P<0.01. D, Diagram of how the Twin-Spot MARCM tool works. E, Representative confocal image of ALG clones generated with Twin-Spot MARCM. F, Number of VNCs with/without clones in controls and Injured VNCs in young (1 day old) and mature (7 day old) animals. Chi-square binomial test, **P<0.01, ****P<0.0001. G-H, Representative confocal images of a two-color ALG clone (G) and a single-color ALG clone (H). I, Percentage of ALG single-color clones in injured controls and Injured VNCs where apoptosis was inhibited in ALG. Unpaired t-test **P<0.01. Genotypes: wild type (B); AlrmGal4>UAS-Fly-FUCCI and R56F03Gal4>UAS-Fly-FUCCI (C); HsFLP; UAS-GFP, UAS-RFPRNAi /UAS-RFP, UAS-GFPRNAi; tubGal80ts:alrmGal4; (E-I): HsFLP; UAS-GFP, UAS-RFPRNAi /UAS-RFP, UAS-GFPRNAi; tubGal80ts:alrmGal4/UAS-p35; (I). Scale bars: 50 μm (B, E); 15 μm (G,H)

Revised version: A, Schematic representation of Drosophila adult Central Nervous System (CNS) and Ventral Nerve Cord (VNC), indicating the MtN, where the injury is performed. B, Representative confocal image showing dividing cells (EdU+) with ALG identity (Repo+ and Pros+) and glial identity different from ALG (Repo+, Pros-). C, Number of ALG or EG cells in G2-M per MtN, in normal conditions and 24 h after injury. Paired t-test, *P<0.05, **P<0.01. C’-C’’, Representative images of ALG in G2-M showing their position within the injured area (C’) and a magnification showing the detail of GFP+RFP+ cells (C’’). D, Diagram of how the Twin-Spot MARCM tool works. D’, Representative confocal image of ALG clones generated with Twin-Spot MARCM. E, Number of VNCs with/without clones in controls and Injured VNCs in young (1 day old) and mature (7 day old) animals. Chi-square binomial test, **P<0.01, ****P<0.0001. F-G, Representative confocal images of a two-color ALG clone (F) and a single-color ALG clone (G). H, Percentage of ALG single-color clones in injured controls and Injured VNCs where apoptosis was inhibited in ALG. Unpaired t-test **P<0.01. Genotypes: wild type (B); AlrmGal4>UAS-Fly-FUCCI (C-C’); R56F03Gal4>UAS-Fly-FUCCI (C); HsFLP; UAS-GFP, UAS-RFPRNAi /UAS-RFP, UAS-GFPRNAi; tubGal80ts:alrmGal4; (D’-H): HsFLP; UAS-GFP, UAS-RFPRNAi /UAS-RFP, UAS-GFPRNAi; tubGal80ts:alrmGal4/UAS-p35 (H). Scale bars: 50 μm (B, C, E); 15 μm (C’’, F,G)

We have studied the physiological mechanisms that occur after CNS injury, focusing on the fate of glial cells upon damage. We provide the first evidence of in vivo physiological reprogramming of glial cells into other glial identities and into new neurons in the adult brain.

Results

Astrocyte-like glia proliferation in adult is increased upon crush injury

Previously we described that glial cells proliferate in basal conditions and in response to crush injury in the adult fly (Losada-Pérez et al., 2021) although we did not identify the glial subtype responsible for this response. Thus, we questioned if those dividing cells were of the NPG type, that is, Astrocyte-Like Glia (ALG), as previous works suggest in larvae (Kato et al., 2020; Losada-Perez et al., 2016), or Ensheathing Glia (EG).

ALG identity is defined by the co-expression of repo and prospero transcription factors (Omoto et al., 2015). To determine if ALG cells divide upon injury, we caused a crush in the metathoracic neuromere of the adult ventral nerve cord (VNC, Figure 1A) (Losada-Pérez et al., 2021) and performed EdU experiments combined with staining with anti-Repo (pan-glia marker) and anti-Prospero antibodies. The results show that ALG cells (Repo positive, Prospero positive) divide (EdU positive) after injury (Figure 1B). Besides, we also observed EdU+ Repo+ cells that were Pros negative. Therefore, this result indicates that glial cells, different from ALG, also undergo division.

There is no specific nuclear marker for EG cells, consequently, to determine if the other dividing cells were EG, we expressed FLY-FUCCI (Zielke et al., 2014) under control of a specific driver for EG (R56F03-Gal4). This genetic tool allows the visualisation of cell-cycle stages by combinations of fluorescent protein reporters (see Materials and Methods). In parallel, we also expressed FLY-FUCCI under control of the ALG specific driver alrm-Gal4 (Altenhein et al., 2006; Beckervordersandforth et al., 2008).

Earlier version: We quantified the number of ALG or EG cells in G2-M (i.e., GFP and RFP positive) in control and injured animals at 6h and 24h after injury. The data show that EG cells have a basal division rate which does not change in response to injury. In contrast, ALG cells have a higher division rate in basal conditions compared to EG, and this rate increases in response to injury (Figure 1C). This result indicates that adult ALG cells activate a proliferative response after injury, akin to the proliferative response in larvae (Losada-Perez et al., 2016). On the other hand, the basal proliferation of EG explains the presence of Repo+EdU+Pros-cells.

Revised version: We quantified the number of ALG or EG cells in G2-M (i.e., GFP and RFP positive) in control and injured animals at 24h after injury. The data show that EG cells have a basal division rate which does not change in response to injury. In contrast, ALG cells have a higher division rate in basal conditions compared to EG, and this rate increases in response to injury (Figure 1C-C’). This result indicates that adult ALG cells activate a proliferative response after injury, akin to the proliferative response in larvae (Losada-Perez et al., 2016). On the other hand, the basal proliferation of EG explains the presence of Repo+EdU+Pros-cells.

To further confirm these results, we generated Twin-MARCM clones (Yu et al., 2009) using specific drivers for EG (R65F03-Gal4) or ALG (alrm-Gal4). We combined the Twin-MARCM technique with tubGal80TS under temperature control (see Materials and Methods). This technique allowed the study of ALG or EG division selectively in adulthood (Figure 1D).

Earlier version: The results of ALG MARCM clones indicate that ALG divides in normal conditions in the adult CNS. Since ALG cells are very heterogeneous in size and shape we were not able to determine the number of cells that composed each clone (Figure 1E). However, we quantified the number of individuals with clones. The results indicate that the percentage of samples with clones is significantly higher in injured animals compared to controls (Figure 1F). To confirm that this division is not the result of residual development, we repeated the experiment with mature (seven-days-old) flies. In this case, the percentage of not-injured CNS with clones is lower compared with young flies (Figure 1G) indicating we were observing a developmental effect. However, in mature flies, the percentage of samples with clones is significantly higher in injured animals compared to not-injured controls (Figure 1G) indicating that ALG divide upon damage after the developmental program is finished.

Revised version: The results of ALG MARCM clones indicate that ALG divides in normal conditions in the adult CNS. Since ALG cells are very heterogeneous in size and shape we were not able to determine the number of cells that composed each clone (Figure 1D’). However, we quantified the number of individuals with clones. The results indicate that the percentage of samples with clones is significantly higher in injured animals compared to controls (Figure 1E). To confirm that this division is not the result of residual development, we repeated the experiment with mature (seven-days-old) flies. In this case, the percentage of not-injured CNS with clones is lower compared with young flies (Figure 1E) indicating we were observing a developmental effect. However, in mature flies, the percentage of samples with clones is significantly higher in injured animals compared to not-injured controls (Figure 1E) indicating that ALG divide upon damage after the developmental program is finished.

MARCM clones induction with the specific driver for EG did not generate any clone, neither in normal conditions nor after crush injury. MARCM technique generates clones in a stochastic manner, therefore the lack of MARCM clones does not indicate that EG never divides, instead, it indicates that this event is less frequent than ALG division. These results are consistent with the FLY-FUCCI results.

Earlier version: Altogether, the results indicate that there is a basal division rate of neuropil glia where EG divides less often than ALG. Also, ALG shows a proliferative response to crush injury. These findings go in line with previous reports indicating the proliferative response of glial cells (Kato et al., 2018; Losada-Perez, 2018). Moreover, the results now describe proliferation of adult fully-differentiated glial cells and define a differential response of EG and ALG in basal conditions, and upon injury. OPCs carry the proliferative response upon injury in mammals (Kato et al., 2018; Losada-Perez, 2018). Therefore, this results further support the homology between ALG and vertebrate OPC.

Revised version: Altogether, the results indicate that there is a basal division rate of neuropil glia where EG divides less often than ALG. Also, ALG shows a proliferative response to crush injury. These findings go in line with previous reports indicating the proliferative response of glial cells (Kato et al., 2018; Losada-Perez, 2018). Moreover, the results now describe proliferation of adult fully-differentiated glial cells and define a differential response of EG and ALG in basal conditions, and upon injury. OPCs carry the proliferative response upon injury in mammals (Kato et al., 2018; Losada-Perez, 2018). Therefore, these results further support the homology between ALG and vertebrate OPC.

Programmed cell dead regulates ALG number after injury

Earlier version: The Twin-MARCM technique implies that the cell division includes homologous recombination. The first division would give rise to two daughter cells, one GFP positive and other RFP positive (Figure 1D). These differentially marked daughter cells could further divide or not, yielding clones of different sizes (Figure 1E and G). However, during ALG MARCM clone quantification we noticed that >50% of clones were single-colour clones (Figure 1 H-I). This means that after the first division, one of the daughter cells either dies or changes its identity (as we use the alrm-Gal4 ALG specific driver to express GFP or RFP). To determine if single colour clones could be explained by Programmed Cell Dead (PCD) we inhibited it by co-expressing UAS-p35 (Hay et al., 1994), while inducing Twin-MARCM clones in injured animals. The results showed that p35 significantly reduced the number of single colour clones (Figure 1I). These results indicate that a portion of dividing ALG cells undergo PCD, reducing the final number of newly generated ALG after injury. This cellular behaviour is commonly observed during development, where cells are produced in excess and are eliminated by PCD later (Fuchs and Steller, 2011; Gabilondo et al., 2018; Miguel-Aliaga and Thor, 2009). Now we describe how differentiated glial cells reproduce this proliferative response, and the elimination of cells by PCD.

Revised version: The Twin-MARCM technique implies that the cell division includes homologous recombination. The first division would give rise to two daughter cells, one GFP positive and other RFP positive (Figure 1D). These differentially marked daughter cells could further divide or not, yielding clones of different sizes (Figure 1D’ and F). However, during ALG MARCM clone quantification we noticed that >50% of clones were single-colour clones (Figure 1 G-H). This means that after the first division, one of the daughter cells either dies or changes its identity (as we use the alrm-Gal4 ALG specific driver to express GFP or RFP). To determine if single colour clones could be explained by Programmed Cell Dead (PCD) we inhibited it by co-expressing UAS-p35 (Hay et al., 1994), while inducing Twin-MARCM clones in injured animals. The results showed that p35 significantly reduced the number of single colour clones (Figure 1H). These results indicate that a portion of dividing ALG cells undergo PCD, reducing the final number of newly generated ALG after injury. This cellular behaviour is commonly observed during development, where cells are produced in excess and are eliminated by PCD later (Fuchs and Steller, 2011; Gabilondo et al., 2018; Miguel-Aliaga and Thor, 2009). Now we describe how differentiated glial cells reproduce this proliferative response, and the elimination of cells by PCD.

NP glia exchange identities between EG and ALG

Apoptosis inhibition significantly reduces ALG single-colour clones, but we could still observe them. This could be explained by the different effectiveness of the UAS-p35 to suppress apoptosis, but also by the existence of dedifferentiation or transdifferentiation processes. To test this hypothesis, we combined the G-TRACE tool, which allows the visualisation of historical and current Gal4 activation, with TubGal80ts to restrict the G-trace historical visualisation to adulthood.

We drove the expression of G-TRACE reporter with the specific drivers for ALG or EG and analysed the historical and current activation of ALG or EG drivers. The results show that in both ALG and EG experiments there were cells with historical driver activation but not current activation (GFP+RFP-) (Figure 2A-A’). Also, those cells have distinctive nuclear sizes and positions. We named the first type neuropil-like cells (NP-like, arrowheads Figure 2A), as they have similar nuclear shape and size to ALG and EG and they were found in the cortex-neuropil interface, where neuropil glia nuclei are usually positioned. The second type of cells were found in the ventral cortex and presented larger nuclei compared to NP-like cells, thus, given their position we named those cells as Ventral Cortex cells or VC cells (arrowheads Figure 2A’).

Neuropil glia transdifferentiate to two different cell types

A-B. Representative confocal slices showing transdifferentiated EG cells with the G-TRACE tool. A, Medial slice showing NP-like cells position within the MtN. B, Zoom slice (merge and separated channels) showing NP-like cells (GPF+, RFP-) have glial identity (repo+). A’, Ventral slice showing VC cells position within the MtN. B’, Zoom slice (merge and separated channels) showing VC cells (GPF+, RFP-) are not glia (repo-). C, Representative ventral slice showing transdifferentiated CV cells visualised with G-TRACE toll driven by the pan-glial factor Repo-Gal4. D, Number of NP-like cells originated from ALG or EG in not injured and injured MtM. Mann Whitney test ns p>0.05, **P<0.01. E, Percentage of NP-like cells compared with the total number of NP glia that originate them. Scale bars: 50 μm (A, A’, C); 15 μm (B,B’). Genotypes: tubGal80ts, R56F03Gal4>UAS G-TRACE (A, B, D, E); tubGal80ts, repoGal4>UAS G-TRACE (C); tubGal80ts, AlrmGal4>UAS G-TRACE (D, E).

Next, we asked if those cells that changed their fate still had a glial identity. To answer this question, we stained NP-like cells with the pan-glial marker Repo. The results indicate that NP-like cells are glia as they were Repo+ (Figure 2 B-B’). To confirm this, we expressed G-TRACE under the control of Repo Gal4 pan-glial driver and analysed the presence or absence of transformed cells. We identified GFP+RFP-cells localised in the ventral cortex with identical nuclear shape as the VC cells identified in previous experiments (Figure 2C), however, we could not find any NP-like cells. Thus, it seems that NP-like cells have glial identity while VC cells have not.

Earlier version: Next, we addressed if the NP-like glia fate change was triggered by injury. To that end we quantified the number of NP-like cells transformed from ALG or from EG, in not-injured and injured animals 6h and 24h after crush injury. The results show that the number of NP-like cells does not increase in response to injury no matter their initial identity, ALG or EG (Figure 2 D). However, we found significant differences in the number of transformed cells that originated from ALG or EG. We observed that ALG to NP-like transformation is less frequent than EG to NP-like transformation (<0.5% in ALG compared with >10% in EG) (Figure 2E). The transformation of EG into ALG has been shown during metamorphosis (Kato et al., 2020), indicating the versatility of those cells. However, this is the first report of EG to ALG transformation in vivo in adult flies. The functional implications of this process include CNS regeneration, cellular homeostasis and behavioural consequences including locomotion recovery, but probably also learning and memory.

Revised version: Next, we addressed if the NP-like glia fate change was triggered by injury. To that end we quantified the number of NP-like cells transformed from ALG or from EG, in not-injured and injured animals 6h and 24h after crush injury. The results show that the number of NP-like cells does not increase in response to injury no matter their initial identity, ALG or EG (Figure 2 D). However, we found significant differences in the number of transformed cells that originated from ALG or EG. We observed that ALG to NP-like transformation is less frequent than EG to NP-like transformation (<0.5% in ALG compared with >10% in EG) (Figure 2E).

Earlier version: Next, we use anti-prospero to determine if the NP-like cells originated from EG have ALG identity (Kato et al., 2020). The images show that the observed NP-like cells originated from EG where pros+ (Figure 3A), suggesting that EG transdifferentiate into ALG. Prospero is a marker of ALG cells thus, to determine if Prospero is required for EG to ALG transformation we quantified and compared the number of NP-like cells with EG origin in controls and in individuals where pros was downregulated specifically in EG cells. The results showed a significant decrease in the number of NP-like cells (Figure 3B). Thus, we can conclude that Prospero is required for the EG to ALG transformation in adult CNS.

Revised version: The transformation of EG into ALG has been shown during metamorphosis (Kato et al., 2020), indicating the versatility of those cells. Thus, we use anti-prospero to determine if the NP-like cells originated from EG have ALG identity in this context. The images show that the NP-like cells originated from EG where pros+ (Figure 3A), suggesting that EG transdifferentiate into ALG. Next, to determine if Prospero is required for EG to ALG transformation we quantified and compared the number of NP-like cells that arose from EG origin in controls and in individuals where pros was downregulated specifically in EG cells. The results showed a significant decrease in the number of NP-like cells upon pros knockdown in EG cells (Figure 3B). These results indicate that pros downregulation in EG reduces inter-glia conversion. Thus, we can conclude that Prospero is required for the EG to ALG transformation in adult CNS.

EG transdifferentiation into ALG requires prospero

A. Representative image of NP-like cells originated from EG stained with ALG marker Prospero. Top left: Merge image with DAPI. Top right: Merge image without DAPI. Bottom left: GFP signal. Bottom right: RFP signal and pros staining. B, Number of NP-like cells originated from EG in controls animals where pros was inhibited in EG. Unpaired t-test *P<0.05. Scale bar: 15 μm (A). Genotypes: tubGal80ts, R56F03Gal4>UAS G-TRACE (A, B)

Added in the revised version: This is the first report of EG to ALG transformation in vivo in adult flies. The functional implications of this process include CNS regeneration, cellular homeostasis and behavioural consequences including locomotion recovery, but probably also learning and memory.

NP-Glia transdifferentiate to neurons upon injury

Earlier version: Repo-GTRACE experiments demonstrated that VC cells are not glia (Figure 2B-C). To test if these new cells have a neuronal identity, we stained the VNCs with the pan-neuronal marker Elav. The analysis showed that these cells were Elav positive (Figure 4A-A). This result indicates that glial cells can transdifferentiate into neurons under normal conditions. Next, we quantified the number of VC cells following injury. The results show that the number of VC cells is significantly higher in injured animals compared to non-injured controls (Figure 4B). Next, we analysed which NPG subpopulation (i.e., ALG or EG) were responsible for this increase of VC neurons. To eliminate any developmental effect, we analysed in mature flies (see M&M) the possible transdifferentiation from ALG or EG into VC in control and injured animals. The result showed that the new neurons originate from both transdifferentiated ALG and EG (Figure 4C).

Revised version: Repo-GTRACE experiments demonstrated that VC cells are not glia (Figure 2B-C). To test if these new cells have a neuronal identity, we stained the VNCs with the pan-neuronal marker Elav. The analysis showed that these cells were Elav positive (Figure 4A-A). This result indicates that glial cells can transdifferentiate into neurons under normal conditions. Next, we quantified the number of VC cells following injury. The results show that the number of VC cells is significantly higher in injured animals compared to non-injured controls (Figure 4B). Next, we analysed which NPG subpopulation (i.e., ALG or EG) were responsible for this increase of VC neurons. To eliminate any developmental effect, we analysed in mature flies (see M&M) the possible trans differentiation from ALG or EG into VC in control and injured animals. The result showed that the new neurons originate from both transdifferentiated ALG and EG (Figure 4C).

EG and ALG transdifferentiate into neurons upon injury through two different mechanisms

A-A’.Transdifferentiated VC cells (GFP+,RFP-) raised from ALG (A) and EG (A’) are elav+. From left to right: Merge with DAPI, merge without DAPI and green and red channels merged. B, Number of transdifferentiated glia (VC cells, GFP+, RFP-) in controls and injured VNCs. STATS. C, Number of transdifferentiated VC cells (GFP+, RFP-) originates from ALG or EG in controls and injured mature animals. D, Number of ALG or EG that remain ALG or EG in controls and injured VNCs. Statistics (B, C, D): Paired t-test ***p<0.001, **p<0.01, ns p>0.05. Scale bar: 15 μm (A). Genotypes: tubGal80ts, R56F03Gal4>UAS G-TRACE and tubGal80ts, AlrmGal4>UAS G-TRACE

Earlier version: So far, we have seen that ALG cells divide upon injury while EG do not. On the other hand, VC neurons increase out from ALG and EG types in response to injury. This could indicate that ALG divides prior to transdifferentiation while EG transdifferentiate directly. To test this hypothesis, we quantified the number of ALG and EG cells in both injured animals and not-injured controls. We found a significant increase of ALG number in injured animals compared to controls in line with our previous results. Moreover, we found a significant decrease of EG in injured animals compared to not-injured controls supporting the direct transformation in response to injury (Figure 4D).

Revised version: So far, we have seen that ALG cells divide upon injury while EG do not. On the other hand, VC neurons increase out from ALG and EG types in response to injury. This could indicate that ALG divides prior to trans differentiation while EG transdifferentiate directly. To test this hypothesis, we quantified the number of ALG and EG cells in both injured animals and not-injured controls. We found a significant increase of ALG number in injured animals compared to controls in line with our previous results. Moreover, we found a significant decrease of EG in injured animals compared to not-injured controls supporting the direct transformation in response to injury (Figure 4D).

Added in the revised version: Finally, to determine the identity of trans differentiated neurons, we stained VC cells with different neuronal markers such GABA, ChaT or GluRIIA (Boppana et al., 2017; Chen et al., 2016; Katheder et al., 2023; Kolodziejczyk et al., 2008). The images showed that VC cells are negative for all (Figure S1), therefore we weren’t able to determine their identity. This result could suggest that these cells are immature.

Taken together, our results indicate that ALG and EG generate new neurons in response to injury through two different mechanisms. ALG divide and transdifferentiate while EG transdifferentiate directly. Thus, glial cells reveal as key contributors to newly generated glia and neurons during regeneration.

Earlier version: Many authors are exploring the potential of astrocytes to generate new neurons, however they are focused on artificial reprogramming or transdifferentiation (Talifu et al., 2022). This is the first description of this phenomenon observed in vivo in wild type animals. Therefore, these findings open a new approach to generate new neurons in a damaged CNS. Understanding the molecular mechanisms that govern this physiological glial behaviour would have an immeasurable value.

Revised version: Many authors are exploring the potential of astrocytes to generate new neurons, however they are focused on artificial reprogramming or trans differentiation (Talifu et al., 2022). This is the first description of this phenomenon observed in vivo in wild type animals. Therefore, these findings open a new approach to generate new neurons in a damaged CNS. Understanding the molecular mechanisms that govern this physiological glial behaviour would have an immeasurable value.

Materials and methods

Fly stocks and genetics

Experiments were performed in adult Drosophila melanogaster flies raised at 25°C or combining 17°C with 29°C for temporal control of UAS/Gal-4 using the thermo-sensitive repression system Gal80 TS. Fly stocks used were: Alrm-Gal4 (BL 67032), R56F03-Gal4 (BL 39157), UAS-GFP.E2f; UAS-mRFP.CycB (Fly-FUCCI, BL 55100), repo-Gal4 (BL-7415), tub Gal80 ts (BL-7019), UAS-p35 (BL5073), UAS-GTRACE (BL 28281), Twin-MARCM stocks (BL56184 and BL56185), UAS-prosRNAi (BL-26745) form Bloomington Drosophila Stock Center and yw hs-FLP from A. Ferrús. To control the activation of different tools we combined the UAS/Gal4 system with the Tubulin-Gal80TS (temperature sensitive) tool. We maintained the desired genotypes at non-permissive temperature (i.e., 17°C, where the Gal80TS repressor protein is active), to block UAS constructs expression.

Crush injury

Crush Injury was performed following the protocol described in Losada-Perez et al. 2021 (Losada-Pérez et al., 2021). When the experiment did not require Gal4/UAS activation exclusively in adult (i.e., UAS Fly-Fucci and Twin-spot MARCM), we breed flies at 25°C; we selected one day old adult flies emerging between 12 to 24h hours before injury; we selected mature flies that emerged 5 to 7 days before injury. In the experiments where we needed to temperature control Gal4 activation (i.e., UAS-GTRACE) we breed flies at 17°C, then we selected that emerged 24 to 48h before injury (for 1-day old flies) or that emerged 10 to 12 days (for mature flies) and then collocated at permissive temperature (i.e., 29° C, where Gal80 TS is inactive and the Gal4/UAS system is active).

Injured flies dragging the third pair of legs that did not show external signs of cuticle damage were selected 24 h later (24 h after crush injury (ACI)).

FLY-FUCCI system

FLY-FUCCI (Zielke et al., 2014) allows the visualization of Cell-cycle activity, GFP or RFP-tagged forms of E2F1 and Cyclin B proteins respectively are expressed using the UAS/Gal4 system. Thus, by the presence/absence of GFP and/RFP it is possible to determine the cell cycle stage. We identify GFP positive, RFP positive cells as cells in the G2-M cell cycle phase.

Twin-spot MARCM (Mosaic Analysis with a Repressible Cell Marker)

Using Twin-spot MARCM technique (Yu et al., 2009) it is possible to detect cells dividing in a specific moment. When a cell divides, each of the two resulting cells is marked with GFP or RFP, always considering that the promoter directing this tool gets activated. This is possible because the genetic constructs used within this tool includes a flipase under the control of a heat shock promoter (HsFlp). This enzyme induces mitotic recombination of sequences situated between FRT (Flipase recognition target) sites. To induce flipase activation, flies were exposed to heat-shocks by introducing the vials in 37ºC water baths for 60 minutes.

Twin-MARCM clone induction

We performed two 60-minute-heat shocks in a water bath at 37 ºC to induce Twin-MARCM clones. Vials with one day old flies received the first heat shock right before crush injury. Then, flies were placed at 25 ºC s for 6 hours. After that, flies received the second heat shock and placed again at 25 ºC. Afterwards, flies were dissected 24h after injury.

G-TRACE tool (Gal4-Technique for Real-time and Clonal Expression)

GTRACE tool (Evans et al., 2009) allows the visualization of historical and current expression of Gal4 gene. It is based on the combination of 3 genetic constructs: UAS-RFP, UAS-FLP and Ubi63-FRT-Stop-FRT-GFP.

Whenever Gal4 expression is activated, it induces the expression of UAS-RFP and UAS-FLP. Flippase enzyme recognizes FRT sites and removes stop sequences, allowing GFP constitutive expression (GFP gene is under the control of a constitutive promoter Ubiquitine-p63E). Cells expressing Gal4, and the posterior lineage, will be marked with GFP. This tool allows tracking the current expression of a Gal4 enhancer sequence (UAS-RFP in red) and its historical expression (Ubi-GFP in green).

G-TRACE activation

Animals that develop at 17°C were selected 24 to 48h after emerging and were collocated at permissive temperature (29° C) for 6 hours. Then, the crush injury was performed. Flies were selected and dissected 24h ACI.

EdU incorporation

We used EdU staining to detect cell proliferation. Flies were injured and left overnight in Eppendorf tubes containing EdU solution (100μM EdU in 5% sugar, 1% dye). 24h after injury flies were dissected in cold 4%FA and fixed 15 minutes in 4%FA while rocking and washed twice with 3%BSA in PBS. Fixed VNCs were permeabilized with PBSTx0,5 (1xPBS, 0,5% Triton X-100) for 20 minutes and washed twice again with 3% Bovine serum albumin in PBS. EdU detection was performed by using the Click-iT® reaction cocktail (Click-iT® EdU, Invitrogen). Afterwards, the CNS were immunostained following the protocol described below.

Immunohistochemistry and EdU detection

Added in the revised version: We performed the immunohistochemistry following the protocol for adult Drosophila immunolabeling according to Purice et al. (2017).

Earlier version: We performed the immunohistochemistry following the protocol for adult Drosophila immunolabeling according to Purice et al. (2017). We used the following antibodies: mouse anti-Repo (DSHB, 8D12 ID: AB_528448, 1:200), mouse anti-Elav (DSHB, 9F8A9 ID: AB_528217), anti-prospero (DSHB, MR1A, ID: AB_528440) guinea pig anti-Repo [gift from B. Altenhein (von Hilchen et al., 2013),

Revised version: We used the following antibodies: mouse anti-Repo (DSHB, 8D12 ID: AB_528448, 1:200), mouse anti-Elav (DSHB, 9F8A9 ID: AB_528217), mouse anti-prospero (DSHB, MR1A, ID: AB_528440), mouse anti-ChAT (DSHB, ID:AB_528122), mouse anti-GluRIIA (DSHB, ID: 528269), guinea pig anti-Repo (gift from B. Altenhein (von Hilchen et al., 2013)), anti-Ph3 (Millipore, 3H10), anti-GABA (Sigma-Aldrich, A2052 (Jackson et al., 1990)).

Quantification, statistical analysis and imaging

Images were acquired by confocal microscopy (LEICA TCS SP5) and processed using Fiji (ImageJ 1.50e) software. Images were assembled using Adobe Photoshop CS4 and Adobe Illustrator CS4. Cells were counted using the cell counter plugin from Fiji or Imaris (Imaris 6.3.1 software-Bitplane). Data were analysed and plotted using GraphPad Prism v7.0.0 or v9.0.0. For qualitative data, analysed in Fig. 1D-D’, we performed a binomial test (**P<0.01). For quantitative data, we used a D’Agostino – Pearson normality test and analysed data with normal distributions by using two-tailed t-test with Welch’s unequal variances t-test. Error bars represent standard deviation (+s.d.).